# Python Basics Exercises

#test

With the weekly exercises, where possible you must fill in the code in the code cells. These are followed by a test cell. Running these cells will return an error if your code is not working. We've added a lot of exercises so you may struggle to complete them all in the session but we wanted to give you practise on most of the bits discussed in the teaching notes. 

## Operators

1. Write a function called `is_even` which will return `True` if the supplied integer is even and `False` if odd

<!-- grade_id: cell-ceadecab14020f80, points: 0, grade: false, solution: true -->
```python
def is_even(number):
    """is_even returns True if number is even and False if odd"""
    # YOUR CODE HERE
    
    return result
```

<!-- grade_id: cell-854920a785760be8, points: 1, grade: true, solution: false -->
```python
# Test the function
assert is_even(5) == False
assert is_even(8) == True
```

2. Write a function called `tens_and_ones` which will take two integers e.g 2 and 3 and combine them to return 23.

<!-- grade_id: cell-09063bd6e2976334, points: 0, grade: false, solution: true -->
```python
def tens_and_ones(tens, ones):
    # YOUR CODE HERE
    
    return result
```

<!-- grade_id: cell-dcd03944734b060f, points: 1, grade: true, solution: false -->
```python
assert tens_and_ones(2,3) == 23
```

3. Write a function called `primes` that takes a single integer and returns a list of all the prime numbers <= to that number

<!-- grade_id: cell-e2bfdebb66dbc4da, points: 0, grade: false, solution: true -->
```python
def primes(n):
    """primes returns a list of all the prime numbers <= n"""
    # YOUR CODE HERE
    
    return prime_numbers

```

<!-- grade_id: cell-2ac7186acd9a45ab, points: 1, grade: true, solution: false -->
```python
# Test the function
assert primes(2) == [2]
assert primes(20) == [2, 3, 5, 7, 11, 13, 17, 19]
```

4. Write a function called `same` that checks if two variables are the same value, and if they are the exact same object in memory. The function should return two boolean values: same_value and same_object

<!-- grade_id: cell-d32e7168cd714944, points: 0, grade: false, solution: true -->
```python
def same(var1, var2):
    # YOUR CODE HERE
    
    return same_value, same_object


```

<!-- grade_id: cell-db7b1103d8c0c1d6, points: 1, grade: true, solution: false -->
```python
var1 = 5
var2 = [5, 5]
var3 = [5, 5]

assert same(var1,var1) == (True, True)
assert same(var1, var2) == (False, False)
assert same(var2, var3) == (True, False)
```

## Slicing lists

5. Given a list of numbers, write a function `slice_a_to_b` which extracts all the elements between index a and index b (inclusive - b should be included) where a and b are arbitrary values. Have the function check that a and b are allowed indices. ie make sure the indices are less than the length of the list. If not the function should return the string 'error'

<!-- grade_id: cell-d34f04a10f898de3, points: 0, grade: false, solution: true -->
```python
def slice_a_to_b(list_of_numbers, a, b):
    # YOUR CODE HERE
    
    return result
```

<!-- grade_id: cell-8e20ced2c62dcb1f, points: 1, grade: true, solution: false -->
```python
numbers = [0,1,2,3,4,5,6,7,8,9]

assert slice_a_to_b(numbers, 2, 6) == [2,3,4,5,6]
assert slice_a_to_b(numbers, 2, 10) == 'error'

```

6. Given a list of strings for example ['apple', 'banana', 'cherry', 'date'], write a function last_two() which extracts the last two elements. Your function should also handle the case that the list only contains 1 or 0 values returning a list with the single value or just None respectively.

<!-- grade_id: cell-ec6d8c6f20734103, points: 0, grade: false, solution: true -->
```python
def last_two(list_of_strings):
    # YOUR CODE HERE
    
    return result
```

<!-- grade_id: cell-e50e9d56b9e0ee20, points: 1, grade: true, solution: false -->
```python
test_strings = ['apple', 'banana', 'cherry', 'date']
test_strings2 = ['just_one']
test_strings3 = []

assert last_two(test_strings) == ['cherry', 'date']
assert last_two(test_strings2) == ['just_one']
assert last_two(test_strings3) == None
```

7. Given a tuple of numbers, write a function called slice_tuple, that extracts the elements from index start to index stop (exclusive - ie so the element at index 4 is not included) with an optional step size that defaults to 2.

<!-- grade_id: cell-824df16f218a4ebc, points: 0, grade: false, solution: true -->
```python
def slice_tuple(numbers, start, stop, step=2):
# YOUR CODE HERE

    return result

```

<!-- grade_id: cell-78c9289a17723c86, points: 1, grade: true, solution: false -->
```python
numbers = (1, 2, 3, 4, 5, 6, 7, 8, 9, 10)
start = 3
stop=7
step=3

assert slice_tuple(numbers, start, stop) == (4, 6)
assert slice_tuple(numbers, start, stop, step=step) == (4, 7)
```

8. Write a function called reverse which **uses slicing** to reverse the order of numbers in a list

<!-- grade_id: cell-46a265c7e2ecb75a, points: 0, grade: false, solution: true -->
```python
def reverse(numbers):
    # YOUR CODE HERE
    
    return result

```

<!-- grade_id: cell-a2fd449bf5894b62, points: 1, grade: true, solution: false -->
```python
numbers = (1, 2, 3, 4, 5, 6, 7, 8, 9, 10)

assert reverse(numbers) == (10, 9, 8, 7, 6, 5, 4, 3, 2, 1)
```

9. Write a function called skip which, given a list of numbers extracts every second number starting at the second number

<!-- grade_id: cell-d53b4e30824a50df, points: 0, grade: false, solution: true -->
```python
def skip(numbers):
    # YOUR CODE HERE
    
    return result

```

<!-- grade_id: cell-ad6d532311d5d362, points: 1, grade: true, solution: false -->
```python
numbers = (1, 2, 3, 4, 5, 6, 7, 8, 9, 10)

assert skip(numbers) == (2, 4, 6, 8, 10)
```

## Working with strings

10. Write a function called `allcaps` that takes a string as input and returns the string in capital letters.

<!-- grade_id: cell-b33057be8e278bc0, points: 0, grade: false, solution: true -->
```python
def allcaps(input_string):
    # YOUR CODE HERE
    
    return result

```

<!-- grade_id: cell-5adf302088795a79, points: 1, grade: true, solution: false -->
```python
# Test the function
input_string = "Hello, World!"
assert allcaps(input_string) == 'HELLO, WORLD!'
```

11. Write a function called `no_space` that removes all the spaces from a sentence

<!-- grade_id: cell-0f976c7c82a3ab1a, points: 0, grade: false, solution: true -->
```python
def no_space(input_string):
    # YOUR CODE HERE
    
    return result

```

<!-- grade_id: cell-e9fcf8cfe1361d8d, points: 1, grade: true, solution: false -->
```python
input_string = 'A string with some gaps         and big gaps'

assert no_space(input_string) == 'Astringwithsomegapsandbiggaps'
```

12. Write a function called `count_character` that counts the number of occurrences of a specific character in a given string.

<!-- grade_id: cell-38a7b10eb39b3c4f, points: 0, grade: false, solution: true -->
```python
def count_character(input_string, character):
    # YOUR CODE HERE
    
    return result

```

<!-- grade_id: cell-d3f4cc5ea9cfd6e3, points: 1, grade: true, solution: false -->
```python
input_string = "Hello, World!"
character = "o"

assert count_character(input_string, character) == 2
```

13. Write a function called `camel_case` which takes a sentence as input and capitalises the first letter of each word. Use the docs as there are lots of handy functions to help you <a href = https://docs.python.org/3/library/stdtypes.html#string-methods> string methods documentation </a>

<!-- grade_id: cell-8ac0f974589bde5b, points: 0, grade: false, solution: true -->
```python
def camel_case(input_sentence):
    # YOUR CODE HERE
    
    return result
```

<!-- grade_id: cell-96ddc2a825eaa963, points: 1, grade: true, solution: false -->
```python
# Test the function
input_sentence = "hello, world! this is a sentence."

assert camel_case(input_sentence) == 'Hello, World! This Is A Sentence.'
```

14. Write a function called `no_vowels` that removes all vowels from a sentence

<!-- grade_id: cell-db3ae6356dc65ee5, points: 0, grade: false, solution: true -->
```python
def no_vowels(input_string):
    # YOUR CODE HERE
    
    return result

```

<!-- grade_id: cell-1ad383767e071ab9, points: 1, grade: true, solution: false -->
```python
# Test the function
input_string = "Hello, World!"

assert no_vowels(input_string) == 'Hll, Wrld!'
```


15. Write a function called `word_count` that counts the number of occurrences of each word in a given sentence and returns a dictionary where the keys are the words and the values are their respective counts.

<!-- grade_id: cell-911bda413c6d305b, points: 0, grade: false, solution: true -->
```python
def count_word_occurrences(input_sentence):
    # YOUR CODE HERE
    
    return word_count


```

<!-- grade_id: cell-098b51cb57314846, points: 1, grade: true, solution: false -->
```python
# Test the function
input_sentence = "This is a test sentence. This is another test sentence."

assert count_word_occurrences(input_sentence) == {'This': 2, 'is': 2, 'a': 1, 'test': 2, 'sentence.': 2, 'another': 1}
```

16. Write a function called `is_anagram` which takes 2 strings and then determines if one string is an anagram (ie the same letters in a different order) of the other. It shouldn't matter whether the letters are lower or uppercase. It should return either `True` or `False`.

<!-- grade_id: cell-cdeaf3f5cb1d5c3d, points: 0, grade: false, solution: true -->
```python
def is_anagram(string1, string2):
    # YOUR CODE HERE
    
    return result


```

<!-- grade_id: cell-e92f922531b511ff, points: 1, grade: true, solution: false -->
```python
# Test the function
string1 = "listen"
string2 = "silent"
string3 = "wrong"

assert is_anagram(string1, string2) == True
assert is_anagram(string1, string3) == False

```

## Functions

17. Write a simple function called `multiply_numbers`, that multiplies two numbers together. It should have one positional argument and one keyword argument with a default value of 4. It should then return the result.

<!-- grade_id: cell-2c354ec24394c074, points: 0, grade: false, solution: true -->
```python
# Write function definition this time
# YOUR CODE HERE

```

<!-- grade_id: cell-c68c14ea8d64334b, points: 1, grade: true, solution: false -->
```python
# Test the function
assert multiply_numbers(2) == 8
assert multiply_numbers(3, 5) == 15
```

18. Write a function called `introduce_myself` that takes two positional arguments. Name and age. When supplied with a name and an age it should return "My name is *name* and I am *age* years old". 

<!-- grade_id: cell-862e7990820535e1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

    


```

<!-- grade_id: cell-624686801f194ecb, points: 1, grade: true, solution: false -->
```python
# Test the function
name = "Alice"
age = 25

assert introduce_myself(name, age) == 'My name is Alice and I am 25 years old.'
```

19. Write a lambda function called `test_lambda` that when given a value x returns 3x^2 - 1

<!-- grade_id: cell-41409b8a0b22989e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-f395ce6df10f08e2, points: 1, grade: true, solution: false -->
```python
assert test_lambda(2) == 11
```

20. Write a function called `even_squares` that uses a list comprehension to create a new list which will just contain the squares of even numbers (odd numbers should be discarded) from a given list.

<!-- grade_id: cell-7807e237d6e43756, points: 0, grade: false, solution: true -->
```python
def even_squares(number_list):
    # YOUR CODE HERE
    
    return result
```

<!-- grade_id: cell-5d90f3238d292261, points: 1, grade: true, solution: false -->
```python
numbers = [2, 7, 4, 9, 6, 5]

assert even_squares(numbers) == [4, 16, 36]
```

21. Write a function called `new_dict` that takes a dictionary. The supplied dictionary will consist of keys which are words and values which are numbers. Use a dict comprehension to return a dictionary with the keys and values swapped. ie the new dictionary will have keys that are numbers and values that are words.

<!-- grade_id: cell-dfa68c18e706e95d, points: 0, grade: false, solution: true -->
```python
def new_dict(dictionary):
    # YOUR CODE HERE
    
    return result

```

<!-- grade_id: cell-a8f04c184b9d6668, points: 1, grade: true, solution: false -->
```python
# Test the function
original_dict = {'apple': 3, 'banana': 5, 'cherry': 2}

assert new_dict(original_dict) == {3: 'apple', 5: 'banana', 2: 'cherry'}

```